---
title: Decision Tree Induction with scikit-learn
math: 
    '\abs': '\left\lvert #1 \right\rvert' 
    '\norm': '\left\lvert #1 \right\rvert' 
    '\Set': '\left\{ #1 \right\}'
    '\mc': '\mathcal{#1}'
    '\M': '\boldsymbol{#1}'
    '\R': '\mathsf{#1}'
    '\RM': '\boldsymbol{\mathsf{#1}}'
    '\op': '\operatorname{#1}'
    '\E': '\op{E}'
    '\d': '\mathrm{\mathstrut d}'
    '\Gini': '\operatorname{Gini}'
    '\Info': '\operatorname{Info}'
    '\Gain': '\operatorname{Gain}'
    '\GainRatio': '\operatorname{GainRatio}'
---

**CS5483 Data Warehousing and Data Mining**
___

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
from sklearn import datasets, tree

%matplotlib widget
plt.close("all")
```

## Decision Tree Induction

We will use the [*iris dataset*](https://en.wikipedia.org/wiki/Iris_flower_data_set) from [`sklearn.datasets` package](https://scikit-learn.org/stable/datasets/index.html)

```python
iris = datasets.load_iris()
```

Recall that the classification task is to train a model that can automatically classify the species (*target*) based on the lengths and widths of the petals and sepals (*input features*).

To build a decision tree, we use `DecisionTreeClassifier` from `sklearn.tree` and apply its `fit` method on the training set.

```python
clf_gini = tree.DecisionTreeClassifier(random_state=0).fit(iris.data, iris.target)
```

To display the decision tree, we can use the function `plot_tree` from `sklearn.tree`:

```python
plt.figure()
tree.plot_tree(clf_gini)
plt.show()
```

To make the decision tree looks better, we can provide additional options:

```python
options = {
    "feature_names": iris.feature_names,
    "class_names": iris.target_names,
    "label": "root",
    "filled": True,
    "node_ids": True,
    "proportion": True,
    "rounded": True,
    "fontsize": 7,
}  # store options as dict for reuse

plt.figure(figsize=(9, 6))
tree.plot_tree(clf_gini, **options)  # ** unpacks dict as keyword arguments
plt.show()
```

In particular, the iris setosa is distinguished immediately after checking the petal width/length.

::::{note}

For each node:
- `___ <= ___` is the splitting criterion for internal nodes, satisfied only by samples going left.
- `gini = ...` shows the impurity index. By default, the algorithm uses the Gini impurity index to find the best binary split. Observe that the index decreases down the tree towards the leaves.
- `value = [_, _, _]` shows the number/fraction of examples for each of the three classes, and `class = ...` indicates the majority class, which may be used as the decision for a leaf node. The majority classes are also color coded. Observe that the color gets lighter towards the root as the class distribution becomes more impure. 

The information of the decision is stored in the `tree_` attribute of the classifier. For more details, run `help(clf_gini.tree_)`.

::::

::::{exercise}
:label: ex:clf_entropy

Assign to `clf_entropy` the decision tree classifier created using *entropy* as the impurity measure. You can do so with the keyword argument `criterion='entropy'` in `DecisionTreeClassifier`. Furthermore, Use `random_state=0` and fit the classifier on the entire iris dataset. Check whether the resulting decision tree is the same as the one created using the Gini impurity index.
::::

<!-- grade_id: tree-entropy, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

plt.figure(figsize=(9, 6))
tree.plot_tree(clf_entropy, **options)
plt.show()
```

<!-- grade_id: same-tree-as-gini, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

::::{caution}

Although one can specify whether to use Gini impurity or entropy, `sklearn` implements neither C4.5 nor CART. In particular, it supports only binary splits on numeric input attributes, unlike C4.5, which supports multi-way splits using the information gain ratio. See a workaround [here][categorical].

[categorical]: https://stackoverflow.com/questions/38108832/passing-categorical-data-to-sklearn-decision-tree

::::

## Splitting Criterion

To induce a good decision tree efficiently, the splitting criterion is chosen 
- greedily to maximize the reduction in impurity and 
- recursively starting from the root.

### Overview using pandas

To have a rough idea of what are good features to split on, we will use [pandas](https://pandas.pydata.org/docs/user_guide/index.html) [`DataFrame`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.html?highlight=dataframe#pandas.DataFrame) 
to operate on the iris dataset.

```python
# write the input features first
df = pd.DataFrame(data=iris.data, columns=iris.feature_names)

# append the target values to the last column
df["target"] = iris.target
df.target = df.target.astype("category").cat.rename_categories(
    dict(zip(range(3), iris.target_names))
)
df
```

To display some statistics of the input features for different classes:

```python
df.groupby("target", observed=False).boxplot(rot=90, layout=(1, 3), figsize=(7, 9))
df.groupby("target", observed=False).agg(["mean", "std"]).round(2)
```

::::{exercise}
:label: ex:good_features

Identify good feature(s) based on the above statistics. Does your choice agree with the decision tree generated by `DecisionTreeClassifier`?

::::

<!-- grade_id: good-features, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Measuring impurity

Suppose nearly all instances of a dataset belong to the same class. In that case, we can return the majority class as the decision without further splitting. A measure of impurity is the Gini impurity index, defined as follows:

::::{prf:definition}

Given a dataset $D$ with a class attribute (discrete target), the Gini impurity index is defined as

$$
\Gini(D):= g(p_0,p_1,\dots)
$$ (Gini)

where $(p_0,p_1,\dots)$ are probability masses corresponding to the empirical class distribution of $D$, and

$$
\begin{align}
g(p_0,p_1,\dots) &:= \sum_k p_k(1-p_k)\\
&= 1- \sum_k p_k^2.
\end{align}
$$ (g)

::::

::::{note}

For convenience, we may also write

- $g(\boldsymbol{p})$ for the stochastic vector $\boldsymbol{p}=\begin{bmatrix}p_0 & p_1 & \dots\end{bmatrix}$ of probability masses, and
- $g(p)$ for the probability mass function $p: k \mapsto p_k$.

::::

We can represent a distribution simply as a numpy array. To return the empirical class distributions of the iris dataset:

```python
def dist(values):
    """Returns the empirical distribution of the given 1D array of values as a
    1D array of probabilities. (The ordering is immaterial.)"""
    counts = np.unique(values, return_counts=True)[-1]
    return counts / counts.sum()


print(f"Distribution of target: {dist(iris.target).round(3)}")
```

The Gini impurity index can be implemented as follows:

```python
def g(p):
    """Returns the Gini impurity of the distribution p."""
    return 1 - (p**2).sum()


print(f"Gini(D) = {g(dist(iris.target)):.3g}")
```

Another measure of impurity uses the information measure called entropy in information theory:

::::{prf:definition}

The information content is defined as 

$$
\Info(D):= h(p_0,p_1,\dots),
$$ (Info)

which is the entropy of the class distribution

$$
\begin{align}
h(\boldsymbol{p}) = h(p_0,p_1,\dots) &= \sum_{k:p_k>0} p_k \log_2 \frac1{p_k}.
\end{align}
$$ (h)

::::

::::{note}

- For convenience, we often omit the constraint $p_k>0$ by regarding $0 \log \frac10$ as the limit $\lim_{p\to 0} p \log \frac1{p} = 0$.
- Unless otherwise specified, all the logarithm is base 2, i.e., $\log = \lg$, where the information quantities are in the unit of bit (binary digit). A popular alternative is to use the natural logarithm, $\log = \ln$, where the unit is in *nat*.

::::

::::{exercise}
:label: ex:h

Complete the following function to compute the entropy of a distribution. You may use the function `log2` from `numpy` to calculate the logarithm base 2.

:::{hint}
:class: dropdown

Consider the solution template:

```python
def h(p):
    ...
    return (p * ___ * ___).sum()
```

:::


::::

<!-- grade_id: entropy, points: 0, grade: false, solution: true -->
```python
def h(p):
    """Returns the entropy of distribution p (1D array)."""
    p = np.array(p)
    p = p[(p > 0) & (p < 1)]  # 0 log 0 = 1 log 1 = 0
    # YOUR CODE HERE
    raise NotImplementedError()


print(f"Info(D): {h(dist(iris.target)):.3g}")
```

<!-- grade_id: test-entropy, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(h([1 / 2, 1 / 2]), 1)
assert np.isclose(h([1, 0]), 0)
assert np.isclose(h([1 / 2, 1 / 4, 1 / 4]), 1.5)
```

<!-- grade_id: htest-entropy, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

### Drop in impurity

::::{prf:definition} 

The drop in Gini impurity for a splitting criterion $A$ on a dataset $D$ with respect to the class attribute is defined as

$$
\begin{align}
\Delta \Gini_{\R{A}}(D) &:= \Gini(D) - \Gini_{\R{A}}(D)\\
\Gini_{\R{A}}(D) &:= \sum_{j} \frac{|D_j|}{|D|} \Gini(D_j),
\end{align}
$$ (Delta-Gini)

where $D$ is split by $A$ into $D_j$ for different outcomes $j$ of the split.

::::

We will consider the binary splitting criterion $\R{X}\leq s$ in particular, which gives

$$
\begin{align}
\Delta \Gini_{\R{A}}(D) = g(\hat{P}_\R{Y}) - \left[\hat{P}[\R{X}\leq s] g(\hat{p}_{\R{Y}|\R{X}\leq s}) + \hat{P}[\R{X}> s]g(\hat{p}_{\R{Y}|\R{X}> s})\right]
\end{align}
$$ (Delta-Gini-binary)

where 

- $\R{Y}$ denotes the target,
- $\hat{P}$ denotes the empirical distribution, and
- $\hat{p}_{\R{Y}|\R{X}\leq s}$, $\hat{p}_{\R{Y}|\R{X}> s}$, and $\hat{p}_{\R{Y}}$ denote the empirical probability mass functions of $\R{Y}$ with or without conditioning.

```python
def drop_in_gini(X, Y, s):
    """Returns the drop in Gini impurity of the target Y
    for the binary splitting criterion X <= s.

    Parameters
    ----------
    X: 1D array
        Input feature values for different instances.
    Y: 1D array
        Target values corresponding to X.
    s: Splitting point for X.
    """
    S = X <= s
    q = S.mean()
    return g(dist(Y)) - q * g(dist(Y[S])) - (1 - q) * g(dist(Y[~S]))


X, Y = df["petal width (cm)"], df.target
print(f"Drop in Gini: {drop_in_gini(X, Y, 0.8):.4g}")
```

To compute the best splitting point for a given input feature, we check every consecutive mid-points of the observed feature values:

```python
def find_best_split_pt(X, Y, gain_function):
    """Return the best splitting points and the maximum gain evaluated using
    gain_function for the split X <= s and target Y.

    Parameters
    ----------
    X: 1D array
        Input feature values for different instances.
    Y: 1D array
        Target values corresponding to X.
    gain_function: function of (X, Y, x)
        A function such as drop_in_gini for evaluating a
        splitting criterion X <= s.

    Returns
    -------
    tuple: (s, g) where s is the best splitting point and g is the maximum gain.

    See also
    --------
    drop_in_gini
    """
    values = np.sort(np.unique(X))
    split_pts = (values[1:] + values[:-1]) / 2
    gain = np.array([gain_function(X, Y, s) for s in split_pts])
    i = np.argmax(gain)
    return split_pts[i], gain[i]


print(
    """Best split point: {0:.3g}
Maximum gain: {1:.3g}""".format(
        *find_best_split_pt(X, Y, drop_in_gini)
    )
)
```

The following ranks the features according to the gains of their best binary splits:

```python
rank_by_gini = pd.DataFrame(
    {
        "feature": feature,
        **(lambda s, g: {"split point": s, "gain": g})(
            *find_best_split_pt(df[feature], df.target, drop_in_gini)
        ),
    }
    for feature in iris.feature_names
).sort_values(by="gain", ascending=False)
rank_by_gini
```

Using the entropy to measure impurity, we have the following alternative gain function:

::::{prf:definition} 

The information gain is defined as 

$$
\begin{align}
\Gain_{\R{A}}(D) &:= \Info(D) - \Info_{\R{A}}(D) && \text{where}\\
\Info_{\R{A}}(D) &:= \sum_{j} \frac{|D_j|}{|D|} \Info(D_j),
\end{align}
$$ (Gain)

::::

We will again consider the binary splitting criterion $\R{X}\leq s$ in particular, which gives

$$
\begin{align}
\Gain_{\R{X}\leq s}(D) = h(\hat{P}_Y) - \left[\hat{P}[\R{X}\leq s] h(\hat{p}_{\R{Y}|\R{X}\leq s}) + \hat{P}[\R{X}> s]h(\hat{p}_{\R{Y}|\R{X}> s})\right]
\end{align}
$$ (Gain-binary)

::::{exercise}
:label: ex:gain

Complete the following function to calculate the information gain on the target $\R{Y}$ for a binary split $\R{X}\leq s$. You may use `dist` and `h` defined previously.

::::

<!-- grade_id: info-gain, points: 0, grade: false, solution: true -->
```python
def gain(X, Y, s):
    """Returns the information Gain of Y for the split X <= s.

    Parameters
    ----------
    X: 1D array
        Input feature values for different instances.
    Y: 1D array
        Target values corresponding to X.
    s: Splitting point for X.
    """
    S = X <= s
    q = S.mean()
    # YOUR CODE HERE
    raise NotImplementedError()


print(f"Information gain: {gain(X, Y, 0.8):.4g}")
```

<!-- grade_id: test-info-gain, points: 1, grade: true, solution: false -->
```python
# tests
rank_by_entropy = pd.DataFrame(
    {
        "feature": feature,
        **(lambda s, g: {"split point": s, "gain": g})(
            *find_best_split_pt(df[feature], df.target, gain)
        ),
    }
    for feature in iris.feature_names
).sort_values(by="gain", ascending=False)
rank_by_entropy
```

<!-- grade_id: htest-info-gain, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

The C4.5 induction algorithm uses information gain ratio instead of information gain:

::::{prf:definition} 

The information gain ratio is defined as 

$$
\begin{align}
\GainRatio_{\R{A}}(D) &:= \frac{\Gain_{\R{A}}(D)}{\operatorname{SplitInfo}_{\R{A}}(D)}
\end{align}
$$ (GainRatio)

which is normalized by 

$$
\begin{align}
\operatorname{SplitInfo}_{\R{A}}(D) &:= h\left(j\mapsto \frac{|D_j|}{|D|} \right)=\sum_j \frac{|D_j|}{|D|}\log \frac{|D|}{|D_j|}.
\end{align}
$$ (SplitInfo)

::::

For binary split $\R{X}\leq s$,

$$
\operatorname{SplitInfo}_{\R{X}\leq s}(D) := h\left(\hat{P}[\R{X}\leq s], \hat{P}[\R{X}> s]\right)
$$ (SplitInfo-binary)

in terms of the empirical distribution.

::::{exercise}
:label: ex:gain_ratio

Complete the following function to calculate the *information gain ratio* for a binary split $\R{X}\leq s$ and target $\R{Y}$.

::::

<!-- grade_id: info-gain-ratio, points: 0, grade: false, solution: true -->
```python
def gain_ratio(X, Y, split_pt):
    S = X <= split_pt
    q = S.mean()
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test-info-gain-ratio, points: 1, grade: true, solution: false -->
```python
# tests
rank_by_gain_ratio = pd.DataFrame(
    {
        "feature": feature,
        **(lambda s, g: {"split point": s, "gain": g})(
            *find_best_split_pt(df[feature], df.target, gain_ratio)
        ),
    }
    for feature in iris.feature_names
).sort_values(by="gain", ascending=False)
rank_by_gain_ratio
```

<!-- grade_id: htest-info-gain-ratio, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

::::{exercise}
:label: ex:difference

Does the information gain ratio give a different ranking of the features? Why?

::::

<!-- grade_id: difference, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE